# Homework 6

## Problem 1

In this problem, you are asked to implement two function. The first functiong gets NFL match data from "Data\nflgames.csv" and the NFL team data from "Data\nflteams.csv", and save them into two dataframes (nflgames_df and nflteams_df).

The second function takes in the nflgames_df and addes two new columns called **"MOV"** and **"HomeWinInd"**. **MOV** is equal to the score won (or lost) by the home team, which is equal to HomeScore - AwayScore. **HomeWinInd** is 1 if the Home team win in the match, 0 otherwise. In other words, **HomeWinInd** is 1 if and only if **MOV** > 1.



<!-- grade_id: cell-b78086c290d4a60c, points: 0, grade: false, solution: true -->
```python
import pandas as pd

#from pandasql import sqldf




def GetData():
    games_df = pd.read_csv("data/nflgames.csv")
    teams_df = pd.read_csv("data/nflteams.csv")    
    return [games_df, teams_df]

GetData()

def AddHomeWin(games_df):

    games_df["MOV"]=games_df["HomeScore"]-games_df["AwayScore"]
    games_df["HomeWinInd"] = games_df["MOV"].apply(lambda x: 1 if x >=1 else 0)
    #for games_df["HomeWinInd"] in range games_df["HomeWinInd"]:
       
    return games_df

```

<!-- grade_id: cell-2853176972493b6e, points: 10, grade: true, solution: false -->
```python
[games_df, teams_df] = GetData()
games_df = AddHomeWin(games_df)
assert games_df.shape == (192,13)
assert teams_df.shape == (32,4)
```

<!-- grade_id: cell-b16dd3056b5a9f95, points: 10, grade: true, solution: false -->
```python
assert round(games_df["MOV"].mean(), 2) == 3.91
assert round(games_df["HomeWinInd"].mean(), 2) == 0.58
```

## Problem 2

In this problem, you are asked to implement a function which groups the games_df dataset by HomeTeamID, and computes the average home win percentage, called **"AvgWinPerc**, and the average home win scores, called **"AvgWinScores"**. The function returns a new data called **hometeam_df**, which contains 3 columns: HomeTeamID, AvgWinPerc and AvgWinScores.

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
#import pyodbc
#import pandas.io.sql as psql
import pandasql as ps

def HomeTeamStat(hometeam_df):
    #hometeam_df = AddHomeWin(games_df)
    hometeam_df = psql("SELECT HomeTeamID, \
    SUM(HomeWinInd)/COUNT(HomeWinInd) AS AvgWinPerc, AVG(HomeScore) AS AvgWinScores FROM games_df WHERE MOV=1 GROUP BY HomeTeamID")
    
  
    return hometeam_df


```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-66-7563495715b4> in <module>()
      1 #import pyodbc
      2 #import pandas.io.sql as psql
----> 3 import pandasql as ps
      4 
      5 def HomeTeamStat(hometeam_df):

ModuleNotFoundError: No module named 'pandasql'
```

<!-- grade_id: cell-681343a4d7e22bbd, points: 15, grade: true, solution: false -->
```python
[games_df, teams_df] = GetData()
games_df = AddHomeWin(games_df)
hometeam_df = HomeTeamStat(games_df)
assert hometeam_df.shape == (32,3)
assert round(hometeam_df["AvgWinPerc"].mean(), 2) == 0.58
assert round(hometeam_df["AvgWinScores"].mean(), 2) == 3.82
```

Output:
```
<ipython-input-67-f0ee7480dd2e> in <module>()
      1 [games_df, teams_df] = GetData()
      2 games_df = AddHomeWin(games_df)
----> 3 hometeam_df = HomeTeamStat(games_df)
      4 assert hometeam_df.shape == (32,3)
      5 assert round(hometeam_df["AvgWinPerc"].mean(), 2) == 0.58

<ipython-input-60-c214df96c44a> in HomeTeamStat(hometeam_df)
      4 def HomeTeamStat(hometeam_df):
      5     #hometeam_df = AddHomeWin(games_df)
----> 6     hometeam_df = psql("SELECT HomeTeamID,     SUM(HomeWinInd)/COUNT(HomeWinInd) AS AvgWinPerc, AVG(HomeScore) AS AvgWinScores FROM games_df WHERE MOV=1 GROUP BY HomeTeamID")
      7 
      8 

TypeError: 'module' object is not callable
```

## Problem 3
In this problem, you are asked to take in the hometeam_df generated by the previous problem, the team_df, and return a new dataset called **"top_hometeam_df"**, which contains 10 top teams ranked by the home win percentage descendingly (and if the home win percetnage ties between 2 teams, then rank the teams by average home win scores). The returned dataframe will contain 10 rows (corresponding to the top 10 teams) and 3 columns: TeamName, AvgWinPerc and AvgWinScore.

Hint: You need to use teams_df to get the nams of the home teams.

<!-- grade_id: cell-a3c8e33b23a52cfd, points: 0, grade: false, solution: true -->
```python
def TopHomeTeam(hometeam_df, teams_df):
    top_hometeam_df = None
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return top_hometeam_df
```

<!-- grade_id: cell-d143e38c322d9286, points: 15, grade: true, solution: false -->
```python
[games_df, teams_df] = GetData()
games_df = AddHomeWin(games_df)
hometeam_df = HomeTeamStat(games_df)
top_hometeam_df = TopHomeTeam(hometeam_df, teams_df)
assert top_hometeam_df["TeamName"].tolist() == \
    ['Green Bay Packers', 'New England Patriots', 'Philadelphia Eagles', 'Denver Broncos', 'Arizona Cardinals', \
     'San Diego Chargers', 'Seattle Seahawks', 'Detroit Lions', 'Indianapolis Colts', 'Baltimore Ravens']
```

## Problem 4 

You are give a dataset which contains the nfl match history in a year. The following figure gives you a description of database. 

<img src="./data/table.png" width="600">

You are asked to write a single query, which takes a team id, and returns a dataframe contains two columns. The first column is opposing_team, which is the team IDs of other teams that have been played against the team with the specified team id. The second column is num_matches, which is the number of times that this opposing team has played against the focal team (specificed by the input). The dataframe should be sorted descendingly with respect to num_matches (**If two teams have the same num_matches, they should be sorted descendingly by their opposing team id**). You are not allowed to use pandas functions except creating a dataframe. You need to use sqldf() and SQL query to get your results.

For example:

    Input: 1,
    Output: 
    
<img src="./data/result.png" height="350" width="200">

<!-- grade_id: cell-72cf7018f4d695a2, points: 0, grade: false, solution: true -->
```python
def opposing_team(team_id):
    nfl_games = pd.read_csv("./data/nflgames.csv")
    df_opposing_team = None
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return df_opposing_team
```

<!-- grade_id: cell-06c7dc215450e8bd, points: 50, grade: true, solution: false -->
```python
assert opposing_team(29).values.tolist() == [[1, 2], [27, 1], [22, 1], [20, 1], [19, 1], \
                                             [15, 1], [13, 1], [12, 1], [11, 1], [7, 1], [6, 1]]

assert opposing_team(1).values.tolist() == [[29, 2], [28, 2], [27, 1], [22, 1], [20, 1], \
                                            [19, 1], [15, 1], [12, 1], [9, 1], [7, 1]]
```